# Assignment No 2c

Text Classification (Sentiment Analysis) Using Bayes Rule 
==============
*Sibt ul Hussain*

## Goal

Your goal in this part of assigment is to implement a Naive Bayes Multinomial classifier using  bag of words model for the classification of text (movie reviews) into different categories..

**Note** Please note that you are allowed to use only those libraries which we have discussed in the class, i.e. numpy, scipy, pandas.

Once you have build and test the model on the provided dataset. You will use the learned techniques to compete in a [Kaggle](https://www.kaggle.com/c/word2vec-nlp-tutorial) competition and report your final score and leaderboard ranking to get full credit.

For final submission attach the screen-shot of the leader-board with your score.

## Submission Instructions
You are required to submit the original notebook file on the Slate (with .ipynb extension), with complete set of outputs. Students failing to do so will get zero marks. 

*Please read each step carefully and understand it fully before proceeding with code writing*

## Plagiarism
Any form of plagiarism will not be tolerated and result in 0 marks.


```python
%pylab inline
import scipy.stats
from collections import defaultdict
```

Output:
```
Populating the interactive namespace from numpy and matplotlib
```

<!-- grade_id: parse_string, points: 0, grade: false, solution: true -->
```python
import re
import bs4

def parse_string(string): 
    """"
        Parse the input string and tokenize it using regular expressisons:
        First clean the string such that it does not have any punctuation or number, it must only have a-z and A-Z.
        Please note that while doing this, the spaces must not get disturbed, but in case of multiple spaces convert 
        them to one space.
        Then convert the string to lower case and return its words as a list of strings.
        
        Example:
        --------
        Input :  computer scien_tist-s are,,,  the  rock__stars of tomorrow_ <cool>  ????
        Output:  ['computer', 'scientists', 'are', 'the', 'rockstars', 'of', 'tomorrow']
        
        Parameters:
        ----------
        string: string to be parsed...
        
        Returns:
        ----------
        list of tokens extracted from the string...
    """
    
    # YOUR CODE HERE
    res = re.sub(r'[^ a-zA-Z]', '', string)
    return res.lower().split()
```

<!-- grade_id: parse_file, points: 0, grade: false, solution: true -->
```python
def parse_file(filename): # Parse a given file
    """"
        Parameters:
        ----------
        filename: name of text file to be read
   
        
        Returns:
        ---------
        read file as raw string (with \n, \t, \r, etc included)
    """
    # YOUR CODE HERE
    with open(filename, 'r') as f:
        out = f.read()
    #raise NotImplementedError()
    return out 
    #return
```

<!-- grade_id: files_to_string, points: 0, grade: false, solution: true -->
```python
def files_to_strings(X):
    
    """
        Read an array (or list) of files where each file content is read in a string...
        Input:
        -------
        X an array (or list) of file names
        
        Returns:
        --------
        OUT as a numpy array with each row containing a read string from the file...
    """
    
    # YOUR CODE HERE
    out = []
    for f in X: 
        res = parse_file(f)
        out.append(res)
    #raise NotImplementedError()
    
    return np.array(out)
```

<!-- grade_id: test_reading_and_parsing, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_list_equal

#assert_list_equal(parse_string("computer scien_tist-s are,,,  the  rock__stars of tomorrow_ <cool>  ????"),
#        [u'computer', u'scientists', u'are', u'the', u'rockstars', u'of', u'tomorrow'], "Incorrect cleanning")


strings = files_to_strings(np.array(["./data/imdb1/neg/cv000_29416.txt", "./data/imdb1/pos/cv000_29590.txt"]))
with open("./data/imdb1/neg/cv000_29416.txt") as f:
    text = f.read()

assert_equal(strings[0], text, "At first index should be text of first file")
assert_equal(strings.shape, (2,), "Shape must be (2,) for two files in list")
```

<!-- grade_id: classifier, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
        
## Hint, you can use python dictionary or default dict for counting the words
# or counter class from collections 

#TODO Complete this class for running the complete classifier... 

#You might need to define auxiliary classes for the complete algorithm..

from collections import defaultdict
from itertools import izip
import math 

def vocabfactory():
    return defaultdict(int)

class NaiveBayes:
    ''' Implements the Naive Bayes For Text Classification... '''
    def __init__(self, classes):
        self.classes=classes
        
        # YOUR CODE HERE
        self.nexamples = 0 
        self.vocab = defaultdict(vocabfactory)
        self.classcounts = defaultdict(int)
        self.wordperclasscounts = defaultdict(int)
        #raise NotImplementedError()
        
    def addExample(self, x, y):
        '''
            Add example to corresponding class model ...
            Input
            ---------
            x: example (list of words)
            y: label...
        '''
        
        # YOUR CODE HERE
        self.classcounts[y] += 1
        self.nexamples += 1 
        words = parse_string(str(x))
        for word in words: 
            self.vocab[word][y] += 1  
            self.wordperclasscounts[y] += 1 
        #raise NotImplementedError()
        
    def train(self, X, Y):
        ''' Train the multiclass (or Binary) Bayes Rule using the given 
            X [m x d] data matrix and Y labels matrix
            
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns:
            -----------
            Nothing
            '''
        # YOUR CODE HERE
        for example, label in izip(X, Y): 
            self.addExample(example, label)
        #raise NotImplementedError()
        
    def test(self, X):
        
        ''' Test the trained classifiers on the given set of examples 
        
                   
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional test examples.
           
            Returns:
            -----------
                pclass: the predicted class for each example, i.e. to which it belongs
        '''
        
        nexamples, nfeatures=X.shape
        # YOUR CODE HERE
        
        out = np.zeros(shape=(nexamples,), dtype='S3')
        for i, example in enumerate(X): 
            out[i] = (self.predict(example))
        #raise NotImplementedError()
        return out 
        
    def predict(self, x):
        '''
            Predict the label of given input example...
            
            Input
            ---------
            x: example (list of words)
            
        '''
        
        # YOUR CODE HERE
        words = parse_string(str(x))
        wordcounts = defaultdict(int)
        vocablen = len(self.vocab.keys())
        
        for w in words: 
            wordcounts[w] += 1 
        
        classprobs = dict()
         
        for k in self.classes: 
            prob = 0.0
            for w in wordcounts.keys(): 
                num = (1.+self.vocab[w][k])
                denum = (vocablen + self.wordperclasscounts[k])
                power = wordcounts[w]
                
                prob += power * (log(num) - log(denum))
            prob += ((1. * self.classcounts[k])/self.nexamples*1.)
            
            classprobs[k] = prob
        #raise NotImplementedError()
 
        return max(classprobs, key=classprobs.get)
        
```

```python
import pandas as pd
import tools as t
```

```python
tdir= "./data/imdb1/" # training dir...
#load data, get list of files for each class...
posfiles=t.get_files(tdir+'/pos','*',withpath=True)
negfiles=t.get_files(tdir+'/neg','*',withpath=True)
```

```python
#generate training and testing data...
plabels=['pos']*len(posfiles)
nlabels=['neg']*len(posfiles)
labels=np.concatenate((plabels,nlabels)) # concatenate the +ve and -ve labels
tX=np.concatenate((posfiles,negfiles))
print "Training data Dimensions =", tX.shape," Training labels dimensions=", labels.shape
```

Output:
```
Training data Dimensions = (2000,)  Training labels dimensions= (2000,)
```

```python
X=files_to_strings(tX) # read files and convert each file into set of strings and return an numpy array
X = X.reshape((X.shape[0], 1))
#Split the data into two halves training and test set...
traindata,trainlabels,testdata,testlabels=t.split_data(X,labels)
#Find the classes to train
classes=np.unique(labels)
```

Output:
```
tools.py:549: VisibleDeprecationWarning: using a non-integer number instead of an integer will result in an error in the future
  traindata.append(Xt[ridx[:ntrainex],:])
tools.py:550: VisibleDeprecationWarning: using a non-integer number instead of an integer will result in an error in the future
  trainlabels.append(Yt[ridx[:ntrainex]])
tools.py:552: VisibleDeprecationWarning: using a non-integer number instead of an integer will result in an error in the future
  testdata.append(Xt[ridx[ntrainex:],:])
tools.py:553: VisibleDeprecationWarning: using a non-integer number instead of an integer will result in an error in the future
  testlabels.append(Yt[ridx[ntrainex:]])
```

```python
#Now build a Naive Bayes classifier and test it...
import time 
starttime = time.time()
print '[Info] training a classifier for following classes {}, {}'.format(classes[0],classes[1])
nb=NaiveBayes(classes) 
nb.train(traindata,trainlabels)
print '[Info] testing a classifier'
pclasses=nb.test(testdata)
elapsedtime = time.time() - starttime
acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
print "[Info] Accuracy = {}".format(acc)    
print '[Info] Elapsed time =', elapsedtime
```

Output:
```
[Info] training a classifier for following classes neg, pos
[Info] testing a classifier
[Info] Accuracy = 0.815
[Info] Elapsed time = 6.91613602638
```

### Test Cells Start
#### Do not Modify

<!-- grade_id: test_classifier_shapes, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_in

nb=NaiveBayes(classes)
nb.train(traindata,trainlabels)
assert_equal (nb.test(testdata).shape[0], testdata.shape[0])
assert_in( type(nb.predict(["ok"])) , [str, np.string_, np.str, np.str_] , "Predict should return a label \
                                                                                            not list or array")
```

<!-- grade_id: test_acc, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_greater

nb=NaiveBayes(classes)
nb.train(traindata,trainlabels)
pclasses=nb.test(testdata)
acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
print 'accuracy =', acc 
assert_greater(acc, 0.77, "Acc must be greater than 77%; you are doing something wrong")    
```

Output:
```
accuracy = 0.815
```

<!-- grade_id: test_classifier_responce, points: 4, grade: true, solution: false -->
```python
from nose.tools import assert_equal

comment_pos = "A nice movie, the case was good. Overall a perfect play"
comment_neg = "A waste of time, cast was bad. a clear No!"

#generate training and testing data...
tX=np.concatenate((posfiles,negfiles))
X=files_to_strings(tX)
X = X.reshape((X.shape[0], 1))

plabels=['pos']*len(posfiles)
nlabels=['neg']*len(posfiles)
true_labels = np.concatenate((plabels,nlabels))
inverted_labels = np.concatenate((nlabels,plabels))

true_nb=NaiveBayes(classes)
true_nb.train(X,true_labels)

inverted_nb=NaiveBayes(classes)
inverted_nb.train(X,inverted_labels)

assert_equal( true_nb.predict(comment_pos.split()), "pos" )
assert_equal( true_nb.predict(comment_neg.split()), "neg" )

assert_equal( inverted_nb.predict(comment_pos.split()), "neg" )
assert_equal( inverted_nb.predict(comment_neg.split()), "pos" )
```

### Test Cells End

# Cross Validation

Now lets throw our methods to winds of different folds and measure their accuracy...

```python
#Now lets generate n-fold training and testing data...
nfolds=10
folds=t.generate_folds(X,labels,nfolds) # generate folds for 
for k in arange(len(folds)):
    print folds[k][0].shape, folds[k][2].shape
```

Output:
```
Generating CV data for 2 classes
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
(1800, 1) (200, 1)
```

```python
totacc=[]
#train a classifier for each fold...
classes=np.unique(labels)

for k in range(nfolds):
    nb=NaiveBayes(classes)
    
    traindata=folds[k][0]
    trainlabels=folds[k][1]
    
    #Lets first train the classifier
    nb.train(traindata,trainlabels)
    
    testdata=folds[k][2]
    testlabels=folds[k][3]
    
    #Lets test the classifier
    pclasses= nb.test(testdata)
    
    #print pclasses
    acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
    print "[Info] Fold {} Accuracy = {}".format(k+1, acc)    
    
    totacc.append(acc)

print totacc

mean_acc = np.mean(totacc)
print '[Info] Mean Accuracy =', mean_acc
```

Output:
```
[Info] Fold 1 Accuracy = 0.825
[Info] Fold 2 Accuracy = 0.835
[Info] Fold 3 Accuracy = 0.8
[Info] Fold 4 Accuracy = 0.825
[Info] Fold 5 Accuracy = 0.8
[Info] Fold 6 Accuracy = 0.825
[Info] Fold 7 Accuracy = 0.78
[Info] Fold 8 Accuracy = 0.825
[Info] Fold 9 Accuracy = 0.845
[Info] Fold 10 Accuracy = 0.785
[0.82499999999999996, 0.83499999999999996, 0.80000000000000004, 0.82499999999999996, 0.80000000000000004, 0.82499999999999996, 0.78000000000000003, 0.82499999999999996, 0.84499999999999997, 0.78500000000000003]
[Info] Mean Accuracy = 0.8145
```

# Excellent, now its time to go into real waters of Kaggle.


You will be needed to create an account on the Kaggle and download the data for the competition ["Bag of words meets bags of popcorn"](https://www.kaggle.com/c/word2vec-nlp-tutorial/data).  Note that you will be only downloading the "labeledTrainData.tsv" and "labeledTestData.tsv".


"labeledTrainData.tsv" will be used for training your model and thus have prespecified labels for each example review. "labeledTestData.tsv" will be used for testing your model and thus don't have prespecified labels for each example. You will predicting the label for each review and then uploading your result to Kaggle server which will be evaluating your model and will give score to your entry. You will report this score during your assignment submission.

**[Caution]** Please note that Kaggle limits maximum number of evaluations per 24 hours to 5 to reduce the overfitting on the test set, so be careful and throughly test your model before submitting your entry to Kaggle server. 

Read the instructions on the Competition Page. Note you are not allowed to use any of the library except what we have learned during class.

```python
# read the data-set
train=pd.read_csv('./data/labeledTrainData.tsv',sep='\t')
```

```python
train.describe()
```

Output:
```
sentiment
count  25000.00000
mean       0.50000
std        0.50001
min        0.00000
25%        0.00000
50%        0.50000
75%        1.00000
max        1.00000
```

```python
train.head()
```

Output:
```
id  sentiment                                             review
0  5814_8          1  With all this stuff going down at the moment w...
1  2381_9          1  \The Classic War of the Worlds\" by Timothy Hi...
2  7759_3          0  The film starts with a manager (Nicholas Bell)...
3  3630_4          0  It must be assumed that those who praised this...
4  9495_8          1  Superbly trashy and wondrously unpretentious 8...
```

```python
Yt=train['sentiment']
Xt=train['review']
Xt=np.array(Xt)
Yt=np.array(Yt)

print Xt.shape
```

Output:
```
(25000,)
```

```python
#read test set...
test=pd.read_csv('./data/testData.tsv',sep='\t')
```

```python
test.head()
```

Output:
```
id                                             review
0  12311_10  Naturally in a film who's main themes are of m...
1    8348_2  This movie is a disaster within a disaster fil...
2    5828_4  All in all, this is a movie for kids. We saw i...
3    7186_2  Afraid of the Dark left me with the impression...
4   12128_7  A very accurate depiction of small time mob li...
```

#### Training Time 

```python
# Let's split the training data into two halves and test our accuracy...
traindata,trainlabels,testdata,testlabels=t.split_data(Xt.reshape((Xt.shape[0],1)),Yt)
classes=np.unique(trainlabels)
```

```python
# Now lets go and train the model and see its performance...
print '[Info] training a classifier for following classes {}, {}'.format(classes[0],classes[1])
nb=NaiveBayes(classes)
starttime = time.time()
nb.train(traindata,trainlabels)
elapsedtime1 = time.time() - starttime
print '[Info] testing a classifier'
starttime = time.time()
pclasses=nb.test(testdata)
elapsedtime2 = time.time() - starttime
pclasses=pclasses.astype(dtype=int) #NOTE: MANUALLY ADDED LINE
acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
print "[Info] Accuracy = {}".format(acc) 
print '[Info] Training time =', elapsedtime1
print '[Info] Testing time =', elapsedtime2
```

Output:
```
[Info] training a classifier for following classes 0, 1
[Info] testing a classifier
[Info] Accuracy = 0.8444
[Info] Training time = 7.4993698597
[Info] Testing time = 101.718343019
```

#### Cross-Validation Time...

```python
pc = pclasses.astype(int)
np.sum(pc == testlabels) / float(testlabels.shape[0])
```

Output:
```
0.84440000000000004
```

```python
#Split the training data into 10 folds and test classifiers performance...

nfolds=10
folds=t.generate_folds(Xt.reshape((Xt.shape[0],1)),Yt,nfolds) # generate folds for 
for k in arange(len(folds)):
    print folds[k][0].shape, folds[k][2].shape
```

Output:
```
Generating CV data for 2 classes
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
(22500, 1) (2500, 1)
```

```python
# As it takes time, so becareful it can cause your machine into red hot oven
totacc=[]
classes=np.unique(Yt)

for k in range(nfolds):
    nb=NaiveBayes(classes)
    
    traindata=folds[k][0]
    trainlabels=folds[k][1]
    
    #Lets first train the classifier
    nb.train(traindata,trainlabels)
    
    testdata=folds[k][2]
    testlabels=folds[k][3]
    
    #Lets test the classifier
    pclasses= nb.test(testdata)
    pclasses= pclasses.astype(int)
    acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
    print "[Info] Fold {} Accuracy = {}".format(k+1, acc)    
    
    totacc.append(acc)

print totacc
print '[Info] Mean Accuracy =', np.mean(totacc)
```

Output:
```
[Info] Fold 1 Accuracy = 0.8464
[Info] Fold 2 Accuracy = 0.8572
[Info] Fold 3 Accuracy = 0.848
[Info] Fold 4 Accuracy = 0.842
[Info] Fold 5 Accuracy = 0.8376
[Info] Fold 6 Accuracy = 0.8548
[Info] Fold 7 Accuracy = 0.8372
[Info] Fold 8 Accuracy = 0.8524
[Info] Fold 9 Accuracy = 0.8368
[Info] Fold 10 Accuracy = 0.848
[0.84640000000000004, 0.85719999999999996, 0.84799999999999998, 0.84199999999999997, 0.83760000000000001, 0.8548, 0.83720000000000006, 0.85240000000000005, 0.83679999999999999, 0.84799999999999998]
[Info] Mean Accuracy = 0.84604
```

# Now let's train on the complete dataset and test on the provided test set...

```python
classes= np.unique(Yt)
print 'Training a Classifier on Full training set with classes =', classes
nb=NaiveBayes(classes)
starttime = time.time()
nb.train(Xt.reshape(Xt.shape[0],1),Yt)
elapsedtime = time.time() - starttime
print 'Elapsed Time =', elapsedtime
```

Output:
```
Training a Classifier on Full training set with classes = [0 1]
Elapsed Time = 10.433978796
```

```python
#Get the test data...
Xtest=test['review']
Xtest=np.array(Xtest.reshape((Xtest.shape[0],1)))
#test the classifier on the provided test set...
starttime = time.time()
pclasses=nb.test(Xtest)
elapsedtime = time.time() - starttime
print 'Elapsed time =', elapsedtime
```

Output:
```
Elapsed time = 441.630002022
```

```python
#write the result in the kaggle's required format
output = pd.DataFrame( data={"id":test["id"], "sentiment":pclasses} )

# Use pandas to write the comma-separated output file
output.to_csv( "./Naive_bays_Bag_of_Words_model.csv", index=False, quoting=3 )
```

# Time to Upload the prediction to Kaggle...

Now upload the result on the Kaggle and see your ranking and score. Using this simple method you can have an accuracy of around 0.80960.

# Improvement by Excluding Stop Words...

You can improve your score further by excluding the commonly occuring words (also known as stop words) in the English language.



```python
#read and create a set of stop 
stopwords=set(t.read_txt_file('./data/english.stop'))
print stopwords
```

Output:
```
set(['secondly', 'all', 'consider', 'whoever', 'four', 'edu', 'go', 'causes', 'seemed', 'rd', 'certainly', 'vs', 'to', 'asking', 'th', 'under', 'sorry', "a's", 'sent', 'far', 'every', 'yourselves', "we'll", 'went', 'did', 'forth', "they've", 'try', 'p', "it'll", "i'll", 'says', "you'd", 'yourself', 'likely', 'further', 'quite', 'even', 'what', 'appear', 'brief', 'goes', 'sup', 'new', 'ever', "c'mon", 'whose', 'respectively', 'never', 'here', 'let', 'others', "hadn't", 'along', "aren't", 'k', 'allows', "i'd", 'howbeit', 'usually', 'whereupon', "i'm", 'changes', 'thats', 'hither', 'via', 'followed', 'merely', 'while', 'viz', 'everybody', 'use', 'from', 'would', 'contains', 'two', 'next', 'few', 'therefore', 'taken', 'themselves', 'thru', 'tell', 'more', 'knows', 'becomes', 'hereby', 'herein', "ain't", 'particular', 'known', 'must', 'me', 'none', 'f', 'this', 'oh', 'anywhere', 'nine', 'can', 'theirs', 'following', 'my', 'example', 'indicated', "didn't", 'indicates', 'something', 'want', '
```

Now you can re-build the model by excluding these words and again upload your results on Kaggle. 

Doing this simple trick can further improve your accuracy to 0.81768.

For final submission attach the screen-shot of the leader-board with your score

Insert ScreenShot of Leader-board Below
----------------------------------------


<img src="./kagglescore.png">